Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

## Lab: basic Python syntax

This lab uses an experimental feature for running the test cells. To be able to run them you have to install a module in the virtual environment

    $ pip install jupyter-pytest
    
It provides more descriptive information when the test does not pass, however, if you happen to make a syntax error  you have to rerun the whole document with 'Kernel -> Restart & Run all'

```python
%load_ext jupytest
```

#### Exercise 1
Write a function 'mymax' that accepts two numbers as parameters and that prints the larger number to the screen. For this elementary exercise, do not use the built-in function 'max()' - just revert to elementary expressions and statements.

<!-- grade_id: cell-ee8c56c7cc588b03, points: 0, grade: false, solution: true -->
```python
def mymax(x, y):
    """
    Returns the maximum of x and y
    """
    # YOUR CODE HERE
    if x > y:
        greater = x
    else:
        greater = y
        
    return greater
    
a=3
b=5

print(f"The maximum of {a} and {b} is {mymax(a, b)}")
```

Output:
```
The maximum of 3 and 5 is 5
```

```python
assert mymax(2, 6) == 6
```

<!-- grade_id: cell-304328f9c7d4c051, points: 0, grade: true, solution: false -->
```python
%%pytest
assert mymax(2, 6) == 6
assert mymax(5, 3) == 5
```

Output:
```
UsageError: Cell magic `%%pytest` not found.
```

<!-- grade_id: cell-304328f9c7d4c051, points: 0, grade: true, solution: false -->
```python
%%pytest
assert mymax(2, 6) == 6
assert mymax(5, 3) == 5
```

Output:
```
UsageError: Cell magic `%%pytest` not found.
```

#### Exercise 2

Write a function to sum the elements of a sequence with a for loop. You can assume that it consists of numeric values.

<!-- grade_id: cell-b607688727ab57f1, points: 0, grade: false, solution: true -->
```python
def mysum(sequence):
    """
    Returns the sum of the elements in sequence
    """
    # YOUR CODE HERE
    s = 0
    
    for i in range(len(sequence)):
        s = s + sequence[i]
        
    for element in sequence:
        s = s + element
        
    return s


print(f'Sum of [1, 2, 3] is {mysum([1, 2, 3])}')
print(f'Sum of (3, 4, 5) is {mysum((3, 4, 5))}')
print(f'Sum of {{}} is {mysum({})}')
```

Output:
```
Sum of [1, 2, 3] is 12
Sum of (3, 4, 5) is 24
Sum of {} is 0
```

<!-- grade_id: cell-24ac435261756312, points: 0, grade: true, solution: false -->
```python
%%pytest
assert mysum([1, 2]) == 3
assert mysum((1, 1/2, 1/4)) == 1.75
```

#### Exercise 3

Variables of `str` type has a method `join` defined such that when a sequence (e.g. `list`, `tuple`) of strings is given as an argument, they are concatenated to a single string with the variable as a separator. E.g.

    >>> separator = '-'
    >>> separator.join(['father', 'in', 'law'])
    'father-in-law'

Use the `join` to generate a histogram with stars such that 

    histogram([3, 4, 1]) 
    
returns

    ***
    ****
    *



<!-- grade_id: cell-631ae09a389a66d4, points: 0, grade: false, solution: true -->
```python
def histogram(sequence):
    """
    Returns histogram in the form of characters
    
    >>> histogram([3, 4, 1])
    ***
    ****
    *
    
    """
    # YOUR CODE HERE
    
    l = []
    for e  in sequence:
        l.append(e * '*')
        print(l)
        
    return '\n'.join(l)

print(f'The histogram of [3, 4, 1] is \n{histogram([3, 4, 1])}')
```

Output:
```
['***']
['***', '****']
['***', '****', '*']
The histogram of [3, 4, 1] is 
***
****
*
```

<!-- grade_id: cell-901ea19debe3ee4c, points: 0, grade: true, solution: false -->
```python
%%pytest
assert histogram([]) == ""
assert histogram([1, 2]) == "*\n**"
assert histogram([1, 2, 3]) == "*\n**\n***"
```

#### Exercise 4

Compute how many times a character is present in a string. Use a dictionary to represent the string. 

<!-- grade_id: cell-94b1ecc182583f9b, points: 0, grade: false, solution: true -->
```python
def character_count(text):
    """
    Returns character count of text
    Input: text of type str
    Output: dict that maps a character to a number count
    """
    # YOUR CODE HERE
    

print(f"Character count of 'ABBA': {character_count('ABBA')}")
```

<!-- grade_id: cell-2a787760db753909, points: 0, grade: true, solution: false -->
```python
%%pytest
assert character_count("ABBA")=={'A': 2, 'B': 2}
assert character_count("CCGGAAGAGCTTACTTAG") == {"A": 5, "C": 4, "G": 5, "T": 4}
```

#### Exercise 5

A simple encryption scheme known as Caesarian roll, is obtained by shifting characters a fixed amount. In Python we can obtain the order of a character with the `ord` function

    >>> ord('A')
    65
    
and there reverse operation

    >>> chr(65)
    'A' 
    
Use this in a function that takes a string and returns the encrypted representation

<!-- grade_id: cell-95b86ffef93462ec, points: 0, grade: false, solution: true -->
```python
def caesar_shift(text, n):
    """
    Returns an encrypted string by shifting input characters by n
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-4403018cd26de943, points: 0, grade: true, solution: false -->
```python
%%pytest
assert caesar_shift('a', 0) == 'a'
assert caesar_shift('ABBA', 2) == 'CDDC'
assert caesar_shift('CDDC', -2) == 'ABBA'
```

Historically, this encryption allegedly used  by the Roman Emperor Caesar (thereby the name). Eventually the encryption was crackedby realizing that some characters are more common than other in a given language. Use the functions defined in previous exercises to

* identify the most common character in the crypto text
* in the English language the most common character is 'e' - what is the shift?
* decrypt the crypto_text with `caesar_shift` using a negative shift

```python
crypto_text = "{v'il3'vy'uv{'{v'ilA'{oh{'pz'{ol'x|lz{pvuA\x11~ol{oly'.{pz'uvisly'pu'{ol'tpuk'{v'z|mmly\x11{ol'zspunz'huk'hyyv~z'vm'v|{yhnlv|z'mvy{|ul3\x11vy'{v'{hrl'hytz'hnhpuz{'h'zlh'vm'{yv|islz3\x11huk'i\x80'vwwvzpun'luk'{oltF'{v'kplA'{v'zsllwB\x11uv'tvylB'huk'i\x80'h'zsllw'{v'zh\x80'~l'luk\x11{ol'olhy{4hjol'huk'{ol'{ov|zhuk'uh{|yhs'zovjrz\x11{oh{'mslzo'pz'olpy'{v3'.{pz'h'jvuz|tth{pvu\x11kl}v|{s\x80'{v'il'~pzo.k5'{v'kpl3'{v'zsllwB\x11{v'zsllwA'wlyjohujl'{v'kylhtA'h\x803'{olyl.z'{ol'y|iB\x11mvy'pu'{oh{'zsllw'vm'klh{o'~oh{'kylhtz'th\x80'jvtl\x11~olu'~l'oh}l'zo|mmslk'vmm'{opz'tvy{hs'jvps3\x11t|z{'np}l'|z'wh|zlA'{olyl.z'{ol'ylzwlj{\x11{oh{'thrlz'jhshtp{\x80'vm'zv'svun'spmlB\x11mvy'~ov'~v|sk'ilhy'{ol'~opwz'huk'zjvyuz'vm'{ptl3\x11{ol'vwwylzzvy.z'~yvun3'{ol'wyv|k'thu.z'jvu{|tls\x803\x11{ol'whunz'vm'klzwpzlk'sv}l3'{ol'sh~.z'klsh\x803\x11{ol'puzvslujl'vm'vmmpjl'huk'{ol'zw|yuz\x11{oh{'wh{plu{'tlyp{'vm'{ol'|u~vy{o\x80'{hrlz3\x11~olu'ol'optzlsm'tpno{'opz'x|pl{|z'thrl\x11~p{o'h'ihyl'ivkrpuF'~ov'~v|sk'mhyklsz'ilhy3\x11{v'ny|u{'huk'z~lh{'|ukly'h'~lhy\x80'spml3\x11i|{'{oh{'{ol'kylhk'vm'zvtl{opun'hm{ly'klh{o3\x11{ol'|ukpzjv}ly.k'jv|u{y\x80'myvt'~ovzl'iv|yu\x11uv'{yh}lssly'yl{|yuz3'w|\x81\x81slz'{ol'~pss\x11huk'thrlz'|z'yh{oly'ilhy'{ovzl'pssz'~l'oh}l\x11{ohu'ms\x80'{v'v{olyz'{oh{'~l'ruv~'uv{'vmF\x11{o|z'jvuzjplujl'kvlz'thrl'jv~hykz'vm'|z'hssB\x11huk'{o|z'{ol'uh{p}l'o|l'vm'ylzvs|{pvu\x11pz'zpjrsplk'v.ly'~p{o'{ol'whsl'jhz{'vm'{ov|no{3\x11huk'lu{lywypzlz'vm'nylh{'wp{o'huk'tvtlu{\x11~p{o'{opz'ylnhyk'{olpy'j|yylu{z'{|yu'h~y\x803\x11huk'svzl'{ol'uhtl'vm'hj{pvu5\x11"
```

<!-- grade_id: cell-440cb014209f1873, points: 0, grade: true, solution: true -->
```python
print("Crypto text:\n", crypto_text)
roll = 0  # redefine this value
# YOUR CODE HERE
print("\nClear text:\n", caesar_shift(crypto_text, -roll))

```

```python

```